Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
##this is used for testing
import numpy as np
import types
from nose.tools import assert_equal, assert_raises
import scipy.special

```

### Binomial Coefficient $N\choose k$ in Python


The above cell imports `scipy.special`, which is a subpackage of `scipy` and contains method `comb(N,k)`, This method can be used to compute binomial coefficient ${N \choose k} = {N\cdot (N-1)\cdot \,.\,.\,.\, \cdot (N-k+1)\over k!} = {N!\over k!\,(N-k)!}$, i.e. the number of ways of choosing a subset of $k$ objects from a set of $N$ objects.

For example, if you need number of ways of choosing 3 out of 7 balls from a box, that number can be found by calling 

`scipy.special.comb(7,3)`


Feel free to use `scipy.special.comb()` in this assignment.

<br>

In addition, if you want to compute factorial $n! = n\cdot (n-1)\cdot (n-2)\cdot\, \dots \, \cdot 2\cdot 1$, you can use 

`scipy.special.factorial(n)`

<br>


### Question 1

Two dice are cast. What is the probability that the sum of the two numbers is at least 10? 

Do NOT round your answer. Be catious! Make sure you count all outcomes.

Write your answer as

`ans = <some expression>`

For example, if you think the answer is `1/36` you should create variable `ans` as

`ans = 1/36`

**IMPORTANT:** Just typing `1/36`, or `answer = 1/36` is not acceptable. Also, do NOT put dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`.


<br>


<!-- grade_id: cell-f6023e5dd46727b2, points: 0, grade: false, solution: true -->
```python

ans = 6/36

```

<!-- grade_id: cell-66b87fb61afca9a5, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-bff80c5cfb3c2b80, points: 7, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 2

Two dice are cast. What is the probability that the sum of the two numbers does not exceed 10 (i.e. is less than or equal to 10)? Do NOT round your answer.

<br>

Write your answer as

`ans = <some expression>`

<br>


<!-- grade_id: cell-51b7dd2c4839fb7d, points: 0, grade: false, solution: true -->
```python

ans = 33/36

```

<!-- grade_id: cell-dd7567e4b92a7299, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-5ba37b4795cab141, points: 7, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 3

What is the probability that in 4 coin tosses we get exactly two heads and two tails? Do NOT round.

<br>

Write your answer as

`ans = <some expression>`

<br>


<!-- grade_id: cell-81782da06f5039bd, points: 0, grade: false, solution: true -->
```python

ans = scipy.special.comb(4,2)/(2**4)

```

<!-- grade_id: cell-40ec419bf1c4739c, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-84c09fbd9b42f199, points: 1, grade: true, solution: false -->
```python
"""Check the accuracy up to rounding to the 1st decimal place"""

assert round(ans,1) == 0.4  , "Sorry, wrong answer"
```

<!-- grade_id: cell-fd92e4accd3c9d4e, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 4

Twenty five elks live in a certain area of which 6 were tagged and released. A certain time later, 5 elks were randomly caught. Assume all 25 elks were equally likely to be caught. What is the probability that exactly 2 of these 5 elks caught were previously tagged? Do NOT round.

<br>

Write your answer as

`ans = <some expression>`

<br>



<!-- grade_id: cell-767bfd46dad418e0, points: 0, grade: false, solution: true -->
```python

ans = (scipy.special.comb(6,2)*scipy.special.comb(19,3))/scipy.special.comb(25,5)

```

<!-- grade_id: cell-30cc2f0253090d26, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-012f5d7c9064e9ef, points: 2, grade: true, solution: false -->
```python
"""Check the accuracy up to rounding to the 1st decimal place"""

assert round(ans,1) == 0.3  , "Sorry, wrong answer"
```

<!-- grade_id: cell-2fe88bfa597b3915, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 5

You pick 2 cards from a regular deck of 52 cards. What is the probability that both of them are aces? Do NOT round your answer.

(Recall that a regular deck of 52 cards consists of 13 kinds/denominations, each having 4 suits: <font color="red">&hearts;, &diams;</font>, &clubs; &spades;.)

<br>

Write your answer as

`ans = <some expression>`

<br>


<!-- grade_id: cell-bbda68dc4b2132bd, points: 0, grade: false, solution: true -->
```python

ans = scipy.special.comb(4,2)/scipy.special.comb(52,2)

```

<!-- grade_id: cell-58245af72aab38ed, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-0445ee461dd2e396, points: 7, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 6

Five cards are dealt from a regular deck of 52 cards. What is the probability of getting a full house: three cards of one kind and two of another kind? Do NOT round your answer. 

(Recall that a regular deck of 52 cards consists of 13 kinds/denominations, each having 4 suits: <font color="red">&hearts;, &diams;</font>, &clubs; &spades;.)

<br>

Write your answer as

`ans = <some expression>`

<br>




<!-- grade_id: cell-a10b3d248b7ec37b, points: 0, grade: false, solution: true -->
```python

ans = (2*scipy.special.comb(13,2)*scipy.special.comb(4,3)*scipy.special.comb(4,2))/scipy.special.comb(52,5)

```

<!-- grade_id: cell-878482512d275b36, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-eb2c651ea089c19e, points: 2, grade: true, solution: false -->
```python
"""Check the accuracy up to rounding to the 1st decimal place"""

assert round(ans,4) == 0.0014   ,  "Sorry, wrong answer"
```

<!-- grade_id: cell-02b027b15987404f, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```