Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
##this is used for testing
import numpy as np
import types
from nose.tools import assert_equal, assert_raises
import scipy.stats, scipy.special

```

### Binomial Coefficient $N\choose k$ in Python


The above cell imports `scipy.special`, which is a subpackage of `scipy` and contains method `comb(N,k)`, This method can be used to compute binomial coefficient ${N \choose k} = {N\cdot (N-1)\cdot \,.\,.\,.\, \cdot (N-k+1)\over k!} = {N!\over k!\,(N-k)!}$, i.e. the number of ways of choosing a subset of $k$ objects from a set of $N$ objects.

For example, if you need number of ways of choosing 3 out of 7 balls from a box, that number can be found by calling 

`scipy.special.comb(7,3)`


Feel free to use `scipy.special.comb()` in this assignment.

<br>

In addition, if you want to compute factorial $n! = n\cdot (n-1)\cdot (n-2)\cdot\, \dots \, \cdot 2\cdot 1$, you can use 

`scipy.special.factorial(n)`

<br>


### Question 1


A die is rolled 10 times. Let $X$ be the number of times six appears in these 10 rolls.



#### Part a)

What is the probability that $X \in \{4,5\}$, i.e. that six appears 4 or 5 times? Do NOT round your answer.


Write your answer as

`ans = <some expression>`

For example, if you think the answer is `1/36` you should create variable `ans` as

`ans = 1/36`

**IMPORTANT:** Just typing `1/36`, or `answer = 1/36` is not acceptable. Also, do NOT put dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`.


<br>


<!-- grade_id: cell-f6023e5dd46727b2, points: 0, grade: false, solution: true -->
```python

ans = (scipy.special.comb(10,4)*((1/6)**4)*((5/6)**6)) + (scipy.special.comb(10,5)*((1/6)**5)*((5/6)**5))

```

<!-- grade_id: cell-66b87fb61afca9a5, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-f975819f6db50750, points: 2, grade: true, solution: false -->
```python
"""Check whether in the 5th and 6th decimal place are 8 and 9, respectively"""

assert np.mod(int(np.floor(ans*10**6)),100) == 89,  "Sorry, wrong answer"

```

<!-- grade_id: cell-bff80c5cfb3c2b80, points: 7, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part b)

What is the probability that six appears even number of times? (i.e. $P(X\in\{0,2,4,6,8,10\}) = ?$). Do NOT round your answer!

Write your answer as 

`ans = <some expression>`

You can use `binom.pmf()` method from `scipy.stats` subpackage to compute probability mass function at desired possible values of random variable $X\sim\mathcal{B}in(n; p)$. To do that, you can call

`scipy.stats.binom.pmf(k, n, p)`

where `k` is an array-like object (a list, range, numpy array, etc.) representing desired values $X$ should take;`n` = number of trials; `p` = probability of success in each trial.

<!-- grade_id: cell-47281b4aeac25572, points: 0, grade: false, solution: true -->
```python

ans = scipy.stats.binom.pmf(0,10,(1/6)) + scipy.stats.binom.pmf(2,10,(1/6)) + scipy.stats.binom.pmf(4,10,(1/6)) + scipy.stats.binom.pmf(6,10,(1/6)) + scipy.stats.binom.pmf(8,10,(1/6)) + scipy.stats.binom.pmf(10,10,(1/6))


```

<!-- grade_id: cell-e56164e03577f41f, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-c9b28e3766fcf718, points: 2, grade: true, solution: false -->
```python
"""Check whether in the 5th and 6th decimal place are 7 and 0, respectively"""

assert np.mod(int(np.floor(ans*10**6)),100) == 70,  "Sorry, wrong answer"
```

<!-- grade_id: cell-02388986b9763bd5, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part c)

What is the expected value $E[X]$? Do NOT round your answer!

Write your answer as 

`ans = <some expression>`

<br>


<!-- grade_id: cell-844c8202fe8a7c47, points: 0, grade: false, solution: true -->
```python

ans = 10 * (1/6)

```

<!-- grade_id: cell-299f2bc522fdb0ba, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer's 7'th decimal digit is 6"""

assert isinstance(ans,float)
assert np.mod(int(np.floor(ans*10**7)),10) == 6,  "Sorry, wrong answer"

```

<!-- grade_id: cell-0c68a28373b1fc7b, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 2

Consider two boxes. Box **A** contains <font style="color: blue">7 blue</font> and <font style="color: red">5 red</font> marbles, while box **B** contains <font style="color: blue">6 blue</font> and <font style="color: red">8 red</font> marbles. A box is selected at random, and 4 marbles are drawn from it at random.

Let $X = $ number of <font style="color: blue">blue </font> balls in the sample of 4 marbles. 


#### Part a)

Create a python function `prob()` with only one input variable `k`, and which for any $k \in \{0,1,2,3,4\}$ computes the value of $P(X=k)$. Do NOT round the output values $P(X=k)$.

<br>




<!-- grade_id: cell-51b7dd2c4839fb7d, points: 0, grade: false, solution: true -->
```python

def prob(k):
    box1_numerator = scipy.special.comb(7,k)*scipy.special.comb(5,4-k)
    box1_denominator = 2*scipy.special.comb(12,4)
    box2_numerator = scipy.special.comb(6,k)*scipy.special.comb(8,4-k)
    box2_denominator = 2*scipy.special.comb(14,4)
    box1 = box1_numerator/box1_denominator
    box2 = box2_numerator/box2_denominator
    return box1+box2

```

<!-- grade_id: cell-dd7567e4b92a7299, points: 1, grade: true, solution: false -->
```python
"""Check whether function prob is defined and has input parameter k"""

assert callable(prob), "Sorry, you didn't define function prob()"
assert (0 <= prob(k=4)) & (prob(k=4) <= 1)  , "Sorry, prob(4) is a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-b0b35c9cfce5ac4a, points: 1, grade: true, solution: false -->
```python
"""Check whether all the probabilities sum up to 1"""

assert abs(sum(np.array([prob(k) for k in range(5)])) - 1) < 1e-8  ,  "Sorry, sum of all prob(k), k=0,1,2,3,4 MUST be 1"
```

<!-- grade_id: cell-ba1b83efc87120e4, points: 1, grade: true, solution: false -->
```python
"""Check whether the 6th decimal digit of prob(k=3) is 7"""

assert np.mod(int(np.floor(prob(3) * 10**6)), 10) == 7 , "Sorry, wrong answer"
```

<!-- grade_id: cell-5ba37b4795cab141, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part b)

Find $E[X]$, the expected number of blue balls in the sample.  Do NOT round your answer!

Write your answer as 

`ans = <some expression>`

<br>

You can use `np.dot()` for dot product of two array-like python objects.

<br>



<!-- grade_id: cell-d84b5d016b3c1a43, points: 0, grade: false, solution: true -->
```python

ans = (0*prob(0)) + (1*prob(1)) + (2*prob(2)) + (3*prob(3)) + (4*prob(4))

```

<!-- grade_id: cell-96d4707c591ce107, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric"""

assert isinstance(ans,float)  ,  "Sorry, wrong answer"

```

<!-- grade_id: cell-c540e957d61275dc, points: 2, grade: true, solution: false -->
```python
"""Check whether the 6th and 7th decimal digits of E[X] are 9 and 5, respectively"""

assert np.mod(int(np.floor(ans * 10**7)), 100) == 95 , "Sorry, wrong answer"
```

<!-- grade_id: cell-09887c75fd5f4e23, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```