Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
##Run this cell first!!! This imports packages used for evaluating test cells
import numpy as np
import types
```

### Question 1.

Which of the following best describes a function whose graph is given below?
Note: The image is also a hyperlink and you can click on it to access an interactive graph.

<a href="image1.html"><img style="border:5px double black;" src="paraboloid1.png" width="300" height="200"></a>

1. $f(x,y) = \sin(x^2+y^2)$

2. $f(x,y) = \sin(x^2+3y^2)$

3. $f(x,y) = x^2 - y^2$

4. $f(x,y) = x^2 - 3y^2$

5. $f(x,y) = x^2 + y^2$

6. $f(x,y) = x^2 + 3y^2$

7. $f(x,y) = {x\over e^y}$

Write your answer as 

`ans = <some number>`

For example, if you think the answer is the 7th of the functions above, your answer should be 

`ans = 7`

**IMPORTANT:** Just typing `7`, or `answer = 7` is not acceptable. Also, do NOT put dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`. 

<!-- grade_id: cell-0a17b28f4e682cc8, points: 0, grade: false, solution: true -->
```python

ans = 6

```

<!-- grade_id: cell-5f58d2a04594b549, points: 2, grade: true, solution: false -->
```python
"""Check whether your answer is an integer b/w 1 and 7"""

assert isinstance(ans,int) , "Sorry, your variable ans is not an integer."
assert (1 <= ans) & (ans <= 7) , "Sorry, your variable ans is not in the range 1-7."
```

<!-- grade_id: cell-f4d035ab0014060a, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 2

Which of the following surfaces represent graphs of fuctions, and which ones don't? 

This image is also a link, so you can click on it to see interactive graphs. Feel free to rotate each plot. When in doubt, look at the various contour lines - they can help you answering the question. 

**Important:** <br>

- There **no** point $(x,y,z)$ of the surface in Plot 4 for which $x = 0$. <br>

- There is **no** point $(x,y,z)$ of the Plot 6 for which $x=y=0$. <br>

- Please read carefully the instructions below the image on how to answer so that the autograder gives you credit!


<a href="six_subplots.html"><img style="border:5px double black;" src="six_subplots_snapshot.png" width="600" height="400"></a>


To answer this question, in the following cell, create a Python list `ans` with string (i.e. character) values "yes" or "no". For example, if you think all plots are graphs of functions except plot 2, you should create a list with 6 components, all being "yes", except the 2nd component, which should be "no". That is, you would type:

`ans = ['yes', 'no', 'yes', 'yes', 'yes', 'yes']`

Of course, this is just an example, it's not the correct answer.

You must have a list of **lowercase** yes'es and no's, written in quotes (you can use either pairs of single or pairs of double quotation marks. Note that the plots are enumerated row-wise. 


<!-- grade_id: cell-633994cf589e968f, points: 0, grade: false, solution: true -->
```python

ans = ['no','no','yes','yes','no','yes']

```

<!-- grade_id: cell-6ea46053badf5b57, points: 1, grade: true, solution: false -->
```python
"""Check that your object ans is a list of strings"""
assert isinstance(ans, list) , "Sorry, your variable ans is NOT a list."
assert set([type(x) for x in ans]) == {str} , "Sorry, all entries of list ans should be strings (i.e. quoted text)"
```

<!-- grade_id: cell-f99221e31baab415, points: 1, grade: true, solution: false -->
```python
"""Check that your object ans is a list of length 6"""
assert len(ans) == 6 , "Sorry, your variable ans is NOT of length 6"
```

<!-- grade_id: cell-fa4eac0f60f2b5a5, points: 1, grade: true, solution: false -->
```python
"""Check that all the components of ans are strings 'yes' or 'no' """
assert set([ans[i].lower() for i in range(len(ans))]) == {"yes","no"} , "Sorry, only values 'yes' and 'no' are allowed."
```

<!-- grade_id: cell-aab080952ff0aa90, points: 7, grade: true, solution: false -->
```python
"""Check whether the list is correct (hidden test)"""

```

Output:
```
'Check whether the list is correct (hidden test)'
```

### Question 3

Find a vector $\vec{n}$ which is orthogonal to the plane $z = 2x - 5y + 3$. Your vector should be defined as a Python's `tuple` object named `n`. Tuples are objects with syntax just like an ordered $n$-tuple in math. For example, if you think the answer is vector $\vec{n} = \langle 1, 2, 3 \rangle$, you should write 

`n = (1, 2, 3)` 

(parentheses make `n` defined as a tuple).


<!-- grade_id: cell-4fbba72269394f0d, points: 0, grade: false, solution: true -->
```python

n = (2,-5,-1)

```

<!-- grade_id: cell-76d36010f9d44ac6, points: 10, grade: true, solution: false -->
```python
"""Check that the variable n is indeed a tuple"""
assert type(n) == tuple


"""Check whether the answer is correct (hidden test)"""
```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 4

Find an equation of the plane that contains point $P=(-2,3,5)$ and is perpendicular to the vector $\vec{n} = \langle 3, -4, 2 \rangle$. Your answer should actually be a Python's lambda expression called `z`, which is of the form

`z = lambda x,y: a * x + b * y + c`

where you just need to find the appropriate values of `a`, `b` and `c`. For example, if you think the solution is the plane $z = 5x - {4\over 5}y - 1$, in the answer cell you should write 


`z = lambda x,y: 5 * x - 4/5 * y - 1`

The key word `lambda` creates a so called lambda expression (a.k.a. anonymous function), which is a function with input variables `x` and `y`. 

<!-- grade_id: cell-cb90e0ee3aaa5bf7, points: 0, grade: false, solution: true -->
```python

z = lambda x,y: -3/2*x + 2*y - 4

```

<!-- grade_id: cell-b0f767966423cc05, points: 2, grade: true, solution: false -->
```python
"""Check whether z is a lambda expression"""
assert isinstance(z,types.LambdaType) , "Sorry, your object z is NOT a lambda expression."


"""Check whether your plane indeed contains the point P"""
assert z(-2,3) == 5 , "Sorry, the graph of your function does NOT contain the point P."
```

<!-- grade_id: cell-5a7095a38eea96c6, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 5

What does the set $\{(x,y,z)\in\mathbf{R}^3\,|\, x = 2y\}$ represent in the 3D space $\mathbf{R}^3$?

1. A line through the origin.

2. A line parallel to the $xy$-plane.

3. A line that does not intersect $z$-axis.

4. A plane that does not intersect $z$-axis.

5. A plane perpendicular to the $xy$-plane.

6. A plane parallel to the $xy$-plane.

7. An airplane.


To answer this question, create an object `ans` and set it equal to the number that corresponds to the correct answer. For example, if you think the correct answer is 7, then type 

`ans = 7`

Hint: It's not 7. &#128515;

<!-- grade_id: cell-b135ef1b089a1f35, points: 0, grade: false, solution: true -->
```python

ans = 5

```

<!-- grade_id: cell-d2d668d11052a753, points: 2, grade: true, solution: false -->
```python
"""Check that your answer is an integer"""
assert isinstance(ans, int)


"""Check that your answer is a number b/w 1 and 7"""
assert (1 <= ans) & (ans <= 7)

```

<!-- grade_id: cell-22078660b609c0e0, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

```python

```