Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
##this is used for testing; run this cell first!
import numpy as np
import types
from nose.tools import assert_equal, assert_raises
```

### Question 1.

Consider the function $\displaystyle f(x,y)={x^2y^3\over x^2+y^4}$. Which of the following is the domain of this function?

<br>

$1. D_f = \{(x,y)\in\mathbf{R}^2\,|\, x\neq 0\}$

$2. D_f = \{(x,y)\in \mathbf{R}^2\,|\, y\neq 0\}$

$3. D_f = \mathbf{R}^2 \setminus \{(0,0)\} = \{(x,y)\in\mathbf{R}^2\,|\,x\neq 0 \text{ or } y\neq 0\}$

$4. D_f = \mathbf{R}$

$5. D_f = \mathbf{R}^2$

$6. D_f = \mathbf{R}^3$

<br>

Write your answer as 

`ans = <some number>`

For example, if you think the answer is the 6th of the choices above, your answer should be 

`ans = 6`

**IMPORTANT:** Just typing `6`, or `answer = 6` is not acceptable. Also, do NOT put a dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`. 

<!-- grade_id: cell-q1-answer, points: 0, grade: false, solution: true -->
```python

ans = 3

```

<!-- grade_id: cell-q1-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether your answer is an integer b/w 1 and 6"""

assert isinstance(ans,int) , "Sorry, your variable ans is not an integer."
assert (1 <= ans) & (ans <= 6) , "Sorry, your variable ans is not in the range 1-6."
```

<!-- grade_id: cell-q1-hidden1, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 2

Consider again the function $\displaystyle f(x,y)={x^2y^3\over x^2+y^4}$.


Find all the values $k\in\mathbf{R}$ for which the limit of $f$ at $(0,0)$ along the path $y=kx$ is 0. That is, find all values of $k\in\mathbf{R}$ such that

$$\lim_{\begin{array}{c|c} (x,y)\to (0,0) \\ y = kx\end{array}}f(x,y)=\lim_{x\to 0}f(x,kx) = 0.$$

<br>

$1.$ For all $k\in \mathbf{R}$.

$2.$ For all $k \neq 0$.

$3.$ For all $k > 0$.

$4.$ For all $k \geq 0$.

$5.$ For no value of $k$ does the limit $\displaystyle \lim_{x\to 0}f(x,kx)$ exist.

$6.$ Only for $k=0$ does the limit exist.

<br>

Write your answer as

`ans = <some number>`

For example, if you think the answer is the 6th of the choices above, your answer should be

`ans = 6`

**IMPORTANT:** Just typing `6`, or `answer = 6` is not acceptable. Also, do NOT put a dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`. 

<!-- grade_id: cell-q2-answer, points: 0, grade: false, solution: true -->
```python

ans = 1

```

<!-- grade_id: cell-q2-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether your answer is an integer b/w 1 and 6"""

assert isinstance(ans,int) , "Sorry, your variable ans is not an integer."
assert (1 <= ans) & (ans <= 6) , "Sorry, your variable ans is not in the range 1-6."
```

<!-- grade_id: cell-q2-hidden1, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 3

Consider the function $\displaystyle f(x,y)={\sin(x+y)\over x+y}$. Which of these is the set of all points $(x,y)$ at which $f$ is continuous?

$1.$ $\{(x,y)\in\mathbf{R}^2\,|\,y\neq -x\}$

$2.$ $\{(x,y)\in\mathbf{R}^2\,|\, x\neq 0 \}$

$3.$ $\{(x,y)\in\mathbf{R}^2\,|\, y\neq 0 \}$

$4.$ $\mathbf{R}^2$

$5.$ $\mathbf{R}^2\setminus \{(0,0)\}$

$6.$ $\emptyset$ &nbsp;  (that is, at no point is the function continuous)

<br>

Write your answer in the following cell as 

`ans = <some number>`

For example, if you think the answer is the 6th of the choices above, your answer should be

`ans = 6`

**IMPORTANT**: Just typing 6, or `answer = 6` is not acceptable. Also, do NOT put a dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called ans, and the code will test the value of ans.

<!-- grade_id: cell-q3-ans, points: 0, grade: false, solution: true -->
```python

ans = 1

```

<!-- grade_id: cell-q3-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether your answer is an integer b/w 1 and 6"""

assert isinstance(ans,int) , "Sorry, your variable ans is not an integer."
assert (1 <= ans) & (ans <= 6) , "Sorry, your variable ans is not in the range 1-6."
```

<!-- grade_id: cell-q3-test2, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 4

Consider the function  $\displaystyle f(x,y) = {x^{m} y^3\over x^2 + y^4}$, where $m\in\mathbf{Z}$ is a fixed integer parameter (we only take integer values to avoid the problem of definability when $x<0$). For which values of $m \in \mathbf{Z}$ does the limit along **any** path $y=kx$ is equal to 0?

In other words, find all $m \in\mathbf{Z}$ for which we have that

for all $k\in \mathbf{R}$, $\displaystyle \lim_{\begin{array}{c|c}(x,y)\to (0,0) \\ y=kx \end{array}}f(x,y) = \lim_{x\to 0}f(x,kx) = 0.$

<br>

$1.$ For $m > 1$

$2.$ For $m > -1$

$3.$ For $m \in \{-1, 0, 1\}$

$4.$ For $m < 1$

$5.$ For $m < -1$

$6.$ For no $m\in \mathbf{Z}$ is the limit at $(0,0)$ along every path $y=kx$ equal to 0.



(**Hint:** As usual, plug in $y=kx$ and try to cancel as high power of $x$ as possible. Also, plotting graphs for various values of $\alpha$ can be very useful, as well. But, keep in mind that sometimes, due to numerical instability or scaling factor, a plot might be misleading and you can only use it for hints and inspiration, not as a proof. Also, the finer the meshgrid, the more accurate the plot is, although that comes at the cost of more computations, and thus, more time.)

<br>

Write your answer in the following cell as 

`ans = <some number>`

For example, if you think the answer is the 6th of the choices above, your answer should be

`ans = 6`

**IMPORTANT**: Just typing 6, or `answer = 6` is not acceptable. Also, do NOT put a dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called ans, and the code will test the value of ans.

<!-- grade_id: cell-q4-ans, points: 0, grade: false, solution: true -->
```python

ans = 4

```

<!-- grade_id: cell-q4-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether your answer is an integer b/w 1 and 6"""

assert isinstance(ans,int) , "Sorry, your variable ans is not an integer."
assert (1 <= ans) & (ans <= 6) , "Sorry, your variable ans is not in the range 1-6."
```

<!-- grade_id: cell-q4-test2, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 5

Find the limit of the function $\displaystyle f(x,y)={x^2 y\over x^2+y^2}$ at point $(0,0)$ along the $y$-axis path: $\{(0,y)\in\mathbf{R}^2\,|\, y\in\mathbf{R}\}$.

<br>

$1.$ $\displaystyle \lim_{\begin{array}{cc}(x,y)\to (0,0)\\ x = 0\end{array}}f(x,y) = {1\over y}$

$2.$ $\displaystyle \lim_{\begin{array}{cc}(x,y)\to (0,0)\\ x = 0\end{array}}f(x,y) = \infty $

$3.$ $\displaystyle \lim_{\begin{array}{cc}(x,y)\to (0,0)\\ x = 0\end{array}}f(x,y) = - \infty$

$4.$ The limit $\displaystyle \lim_{\begin{array}{cc}(x,y)\to (0,0)\\ x = 0\end{array}}f(x,y)$ does not exist, not even in an infinite sense.

$5.$ $\displaystyle \lim_{\begin{array}{cc}(x,y)\to (0,0)\\ x = 0\end{array}}f(x,y) = 0$

$6.$ $\displaystyle \lim_{\begin{array}{cc}(x,y)\to (0,0)\\ x = 0\end{array}}f(x,y) = 1$





<br>

Write your answer as

`ans = <some number>`

For example, if you think the answer is the 6th of the choices above, your answer should be

`ans = 6`

**IMPORTANT:** Just typing `6`, or `answer = 6` is not acceptable. Also, do NOT put a dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`. 

<!-- grade_id: cell-q5-answer, points: 0, grade: false, solution: true -->
```python

ans = 5

```

<!-- grade_id: cell-q5-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether your answer is an integer b/w 1 and 6"""

assert ans in {1,2,3,4,5,6},  "Sorry, variable ans is not an integer between 1 and 6"
```

<!-- grade_id: cell-q5-test2, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 6

Create a python function `squares()` which for any given python list `x` of numeric values  returns the list of squares of components of vector `x`.

For example, if `x = [-2, 3, 0]`, your function should return

`squares(x)`<br>
`[4,9,0]`

If `x = [2, -5, 4, 3, 1, 1, -3]`, your function should return

`squares(x)` <br>
`[4, 25, 16, 9, 1, 1, 9]`

<!-- grade_id: cell-q6-answer, points: 0, grade: false, solution: true -->
```python

def squares(x):
    return [i ** 2 for i in x]

```

<!-- grade_id: cell-q6-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether the output for x = [-2, 3, 0] is a list of length 3"""

x = [-2, 3, 0]
assert isinstance(squares(x),list) , "Sorry, your output is not a list."
assert len(squares(x)) == 3, "Sorry, wrong length of the output list"
```

<!-- grade_id: cell-q6-test2, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is correct for x = [-2, 3, 0]"""

x = [-2, 3, 0]
assert squares(x) == [4, 9, 0]
```

<!-- grade_id: cell-q6-test3, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is correct for x = [2, -5, 4, 3, -1, 1, -3]"""

x = [2, -5, 4, 3, -1, 1, -3]
assert isinstance(squares(x),list) , "Sorry, your output is not a list"
assert squares(x) == [4, 25, 16, 9, 1, 1, 9] 
```

<!-- grade_id: cell-q6-hidden1, points: 4, grade: true, solution: false -->
```python
"""Check whether the answer is correct for various lists x (hidden test)"""

```

Output:
```
'Check whether the answer is correct for various lists x (hidden test)'
```

### Question 7

Create a python function `sumsquares()` which, for a whole number $n\geq 0$, returns $\sum_{k=0}^n k^2$, i.e. the sum of squares of all the whole numbers up to and including $n$. If $n$ is not a whole number (i.e. not a non-negative integer), use the statement <br>

`raise ValueError("Sorry, n must be a whole number!")` <br>

in your code to report the error. Do NOT change/alter the error message, it must be as above.  <br>

 Here are some examples: <br>

`sumsquares(0)` <br>
`0` 

`sumsquares(2)` <br>
`5` 

`sumsquares(2.3)` <br>
`ValueError: Sorry, n must be a whole number!`


<!-- grade_id: cell-q7-answer, points: 0, grade: false, solution: true -->
```python

def sumsquares(n):    
    summation = 0;
    if type(n) != int or n < 0:
        raise ValueError("Sorry, n must be a whole number!")
    else:
        for i in range(1,n+1):
            summation = summation + (i * i)
    return summation

```

<!-- grade_id: cell-q7-test1, points: 1, grade: true, solution: false -->
```python
"""Check whether the function gives correct answer for n=0 and for n=3"""
assert sumsquares(0) == 0, "Sorry, sumsquares(0) should be 0"
assert sumsquares(3) == 14, "Sorry, sumsquares(3) should be 14, since 1^2 + 2^2 + 3^2 = 14"
```

<!-- grade_id: cell-q7-test2, points: 1, grade: true, solution: false -->
```python
"""Check whether the function raises the ValueError if n is not a non-negative integer"""
assert_raises(ValueError, sumsquares, 1.2)
assert_raises(ValueError, sumsquares, -1)

"""Check the exact content of the error message"""
try:
    sumsquares(1.2)
except ValueError as ve:
    err = str(ve)

assert err == "Sorry, n must be a whole number!" , "Your error message is wrong! Should be as suggested (verbatim!)."
```

<!-- grade_id: cell-q7-hidden1, points: 3, grade: true, solution: false -->
```python
"""Check whether the function gives correct answers for various values of n (hidden test)"""

```

Output:
```
'Check whether the function gives correct answers for various values of n (hidden test)'
```

<!-- grade_id: cell-q7-hidden2, points: 3, grade: true, solution: false -->
```python
"""Check whether the function raises the ValueError if n is not a non-negative integer (hidden test)"""

```

Output:
```
'Check whether the function raises the ValueError if n is not a non-negative integer (hidden test)'
```