Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
## DON'T FORGET to run this cell first
##these modules are used for testing; 
import numpy as np
import types
from nose.tools import assert_equal, assert_raises
```

### Question 1.

Consider the function $f(x,y)=x^2\sin^3 y + 2{y^3\over \sqrt{x}}$.

Define Python function `gradf(x,y)` which for each point `(x,y)` with $x>0$ gives the gradient $\nabla f(x,y) = (f_x(x,y), f_y(x,y))$ as Python's tuple object.

Use attribute `pi` from `numpy` for the value of $\pi$. So, if `numpy` is imported as `np`, use `np.pi`

<!-- grade_id: cell-0a17b28f4e682cc8, points: 0, grade: false, solution: true -->
```python

def gradf(x,y):
    dwrtx = (-y**3+2*x**(5/2)*(np.sin(y))**3)/x**(3/2)
    dwrty = 3*x**2*(np.sin(y))**2*np.cos(y)+6*y**2/np.sqrt(x)
    return (dwrtx,dwrty)

```

<!-- grade_id: cell-5f58d2a04594b549, points: 1, grade: true, solution: false -->
```python
"""Check whether gradf is a function with a tuple output"""

assert callable(gradf) , "Sorry, you need gradf defined as a function"
assert isinstance(gradf(2,np.pi/4),tuple) , "Sorry, your output is not a tuple."
assert len(gradf(2,np.pi/4))==2 , "Sorry, the answer should be a tuple of length 2"
```

<!-- grade_id: cell-001ad6f2fced1f85, points: 2, grade: true, solution: false -->
```python
"""Check whether partial derivatives at (2,pi/4) i.e. 1st and 2nd entry of the vector gradf(2,np.pi/4) are correct"""

assert abs(gradf(2,np.pi/4)[0] - 1.24292646471695) <= 1e-5  , "Sorry, partial deriv. w.r.t. x is wrong"

assert abs(gradf(2,np.pi/4)[1] - 6.85971476198386) <= 1e-5  , "Sorry, partial deriv. w.r.t. y is wrong"


```

<!-- grade_id: cell-00dfd9fd0e351ff0, points: 2, grade: true, solution: false -->
```python
"""Check whether gradf(3,np.pi/6) and gradf(5,np.pi/3) are correct"""

assert np.linalg.norm(np.array(gradf(3,np.pi/6))-np.array((0.72237426, 6.79537460))) <= 1e-5  , "Sorry, gradf(5,np.pi/3) is wrong"

assert np.linalg.norm(np.array(gradf(5,np.pi/3))-np.array((6.39247624, 31.0675475))) <= 1e-5  , "Sorry, gradf(5,np.pi/3) is wrong"

```

<!-- grade_id: cell-f4d035ab0014060a, points: 8, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Computing Derivative Numerically

<br>

If $g:x \mapsto g(x)$ is a function of one variable with very complicated formula for derivative, or, if the function (and thus its derivative) is not known at the time of writing the code, we can still write a code to compute the derivative using some numerical approximations. This is also useful when we have a lot of pairs $(x,g(x))$ of data (with $x$'s densely populated), but don't know the exact algebraic expression for $g$. 

The derivative $g'(a)=\displaystyle \lim_{u\to 0}{g(a+u)-g(a)\over u}$ can be computed by using one of the following:

<br>

*forward difference operator*

$\text{DForward}_g(a) = {g(a+h)-g(a)\over h}$, for small step size $h>0$;

<br>

*backward difference operator*

$\text{DBackward}_g(a) = {g(a)-g(a-h)\over h}$, for small $h>0$; 

<br>

*central difference operator*

$\text{DCenter}_g(a) = {\text{DForward}_g(a)\,\, + \,\, \text{DBackward}_g(a) \over 2} = {g(a+h)-g(a-h)\over 2h}$, for small step size $h>0$

<br>

(check that the last formula is indeed the average of the first two).

Based on their form, it should be clear why backward and forward difference operators give us approximations to the derivative $g'(a)=\displaystyle \lim_{u\to 0}{g(a+u)-g(a)\over u}$ &nbsp; (note that for $u=-h$, we have ${g(a)-g(a-h)\over h} = {g(a+u)-g(a)\over u}$).

Therefore, their average, i.e. the central difference operator, is also an approximation of the derivative $g'(a)$. See the graph below to compare the three slopes geometrically they all approximate $g'(a)$ = the slope of the black line. The smaller the $h$, the better the approximations. 

<br>

<img src="approx-derivative.png" style="width: 700px;"/>

<br>

Since $\text{DCenter}_g(a)$ is the average of the forward and the backward difference formulas, it is not surprising that, as it turns out, the central difference formula is more accurate than the first two. More precisely, $\text{DCenter}$ has 2nd order convergence rate, while $\text{DForward}$ and $\text{DBackward}$ are of the 1st order.


<p style="font-size: 16px">In the following problem, we apply centeral difference operator to compute partial derivatives of a function of two variables.</p>

<br>

### Question 2

Consider again the function  $f(x,y)=x^2\sin^3 y + 2{y^3\over \sqrt{x}}$.

In this problem we compute the gradient $\nabla f = (f_x, f_y)$ numerically, using *centeral difference formula*, as explained below.

First, create a Python function `f(x,y)` which, for any point $(x,y)$ with $x>0$, gives the value 

$x^2\sin^3 y + 2{y^3\over \sqrt{x}}$ 

as an output (of type `float`).

<br>

Then, in the same notebook cell, write a Python function `gradfnum(x,y,h)` which for a given input point $(x,y)$ returns the gradient $\nabla f(x,y) = (f_x(x,y), f_y(x,y))$ as a Python `tuple` object, computed numerically with the step size `h`>0, as explained below.

For the step size `h>0`, use the central difference approximations

$f_x(a,b) \approx {f(a+h,b)-f(a-h,b)\over 2 h}$

and

$f_y(a,b) \approx {f(a,b+h) - f(a,b-h) \over 2h}$

<br>

So, the output of `gradfnum()` should be 2-tuple with these two aproximated values. 

<br>

**Note:** Each of the python functions `f(x,y)` and `gradfnum(x,y,h)` could be defined with just one or two lines of the code. You should use `numpy` methods `sin()` and `sqrt()` in `f(x,y)`. As per `gradfnum(x,y,h)`, no methods from any python module should be used in - just use the formulas given above.


<!-- grade_id: cell-633994cf589e968f, points: 0, grade: false, solution: true -->
```python

def f(x,y):
    return x**2*(np.sin(y))**3+2*y**3/np.sqrt(x)

def gradfnum(x,y,h):
    return ((f(x+h,y)-f(x-h,y))/(2*h),(f(x,y+h)-f(x,y-h))/(2*h))

```

<!-- grade_id: cell-758fe293b156800a, points: 2, grade: true, solution: false -->
```python
"""Check some values of f(x,y)"""

assert np.abs(f(4,np.pi)-31.00627668) < 1e-4
assert np.abs(f(1,np.pi/6) - (1/8 + 2*(np.pi/6)**3)) < 1e-4
```

<!-- grade_id: cell-6ea46053badf5b57, points: 2, grade: true, solution: false -->
```python
"""Check some values of mygradf"""

assert np.linalg.norm(np.array(gradfnum(2,np.pi,h=1e-5)) - np.array((-10.96237425, 41.87318520))) <= 1e-4
assert np.linalg.norm(np.array(gradfnum(5,np.pi/3,h=1e-5)) - np.array((6.39247624, 31.0675475))) <= 1e-4
```

<!-- grade_id: cell-f99221e31baab415, points: 9, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 3

Create a python function `grad(f,x,y,h)` which for *any* function $f$ computes its gradient at $(x,y)$ numerically, using the central difference formula (like in the previous problem) and the step size $h>0$. The only difference from the previous problem is that now the function `f` is arbritrary, i.e. is an input argument.

Note: The python function `grad(f,x,y,h)` could be defined with just one or two lines of the code. No python module should be used!

<!-- grade_id: cell-q3-ans, points: 0, grade: false, solution: true -->
```python

def grad(f,x,y,h):
    return ((f(x+h,y)-f(x-h,y))/(2*h),(f(x,y+h)-f(x,y-h))/(2*h))

```

<!-- grade_id: cell-q3-test1, points: 3, grade: true, solution: false -->
```python
"""Check grad for functions x^2*y and x^2+y^2 at some points"""

assert np.linalg.norm(np.array(grad(lambda x,y: x**2 *y, 2,-3, 1e-4)) - np.array((-12,4))) <= 1e-5
assert np.linalg.norm(np.array(grad(lambda x,y: x**2 *y, -2, 5, 1e-4)) - np.array((-20,4))) <= 1e-5
assert np.linalg.norm(np.array(grad(lambda x,y: x**2 + y**2 + 2, 0.4, -0.2, 1e-4)) - np.array((0.8, -0.4))) <= 1e-5
```

<!-- grade_id: cell-q3-test2, points: 9, grade: true, solution: false -->
```python
"""Check whether the answer is correct for some more functions and points (hidden test)"""

```

Output:
```
'Check whether the answer is correct for some more functions and points (hidden test)'
```

### Question 4

Create Python function `DDeriv(f,x,y,h,v)` which, for a given function $f$ and given point $(x,y)$, step size $h>0$ and vector $\vec{v}$, computes directional derivative $Df_{{\vec{v}\over \|\vec{v}\|}}(x,y)$ of $f$ at $(x,y)$ in the direction $\vec{v}$. Note that we need to normalize $\vec{v}$ to compute this directional derivate by

$$D_{{\vec{v}\over \|\vec{v}\|}}f(x,y) = \langle f_x(x,y) , f_y(x,y) \rangle \cdot {\vec{v}\over \|\vec{v}\|}$$

where the partial derivatives are computed numerically by central difference formulae:

$f_x(x,y) \approx {f(x+h,y)-f(x-h,y)\over 2 h}$

$f_y(x,y) \approx {f(x,y+h) - f(x,y-h) \over 2h}$

<br>

You may want to use `numpy` method `dot()` for dot product, and `linalg.norm()` for the norm of a vector.

That is, if `numpy` is imported with alias `np` (as it is at the beginning of this notebook), then 

<br>

* use `np.dot(v1,v2)` for dot product of vectors `v1` and `v2` (with `v1` and `v2` being tuples or list objects)


* use `np.linalg.norm(v)` for the norm of the vector `v` (with `v` being a tuple or a list object)


<!-- grade_id: cell-q4-ans, points: 0, grade: false, solution: true -->
```python

def DDeriv(f,x,y,h,v):
    dwrtx = (f(x+h,y)-f(x-h,y))/(2*h)
    dwrty = (f(x,y+h)-f(x,y-h))/(2*h)
    norm_vec = np.linalg.norm(v)
    return (np.dot((dwrtx,dwrty),v/norm_vec))

```

<!-- grade_id: cell-q4-test1, points: 3, grade: true, solution: false -->
```python
"""Check whether your answer is correct for some functions and at some points and in some directions"""

assert np.abs(DDeriv(f = lambda x,y: x**2+y**2+2, x=0.4, y=-0.2, h=1e-5, v=(0.6,0.8)) - 0.16) <= 1e-4
assert np.abs(DDeriv(f = lambda x,y: x**2 *y - 3* y**3, x=-2, y=1, h=1e-5, v=(3,4)) - (-6.4)) <= 1e-4

```

<!-- grade_id: cell-q4-test2, points: 9, grade: true, solution: false -->
```python
"""Check whether your answer is correct in some more cases (hidden test)"""

```

Output:
```
'Check whether your answer is correct in some more cases (hidden test)'
```