Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
##this is used for testing; run this first
import numpy as np
import math
import types
from nose.tools import assert_equal, assert_raises
import scipy.stats, scipy.special
```


## Question 1

Let $X$ be a random variable with density

$f(x) = \begin{cases}C(4x-x^2) &, 0\leq x \leq 3 \\ 
\hspace{.7cm} 0 &, \text{otherwise} \end{cases}$

where 𝐶∈𝐑 is some constant.

**Draw the graph of this piece-wise density function. This always helps a great deal, when computing values of various atributes of the distribution of $X$. It also helps to spot a possible error in computation, since often you can tell based on the graph of the density function whether certain computed value is plausible or doesn't make any sense.**

<br>


### Part a)

Find the constant $C$.

Write your answer as

`C = <some expression>`

For example, if you think the answer is ${12 \over 34}$ you should create variable C as

`C = 12/34`

**IMPORTANT:** Just typing `12/34`, or `answer = 12/34` is not acceptable. Your answer will be graded automatically, and the test code expects you to create a Python variable called `C`, and the code will test the value of `C`. Do not round your answer.


<!-- grade_id: cell-7a9a738d84913a50, points: 0, grade: false, solution: true -->
```python

C = 1/9

```

<!-- grade_id: cell-235755a990d38276, points: 1, grade: true, solution: false -->
```python
"""Check whether C is numeric (integer or float)"""

assert type(C) in {int, float}  ,  "Sorry, C is not numeric variable"
```

<!-- grade_id: cell-ba5e7e0e4c08d1a0, points: 2, grade: true, solution: false -->
```python
"""Check whether in the 5th and 6th decimal place are 1 and 1, respectively"""

assert np.mod(int(np.floor(C*10**6)),100) == 11,  "Sorry, wrong answer"

```

<!-- grade_id: cell-9c83f1f348fa35fa, points: 7, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```


### Part b)

Find the cumulative distribution function (cdf) of $X$. Create a python function $F(x)$, which for any input $x\in \mathbf{R}$, returns the value of the cdf, i.e. $P(X\leq x)$. Do NOT round the output value.


<!-- grade_id: cell-b594985081bb500a, points: 0, grade: false, solution: true -->
```python

def F(x):
    if x > 3: 
        return 1
    return (1/9)*(2*x**2 - (1/3)*x**3)


```

<!-- grade_id: cell-214b8e292322dbb3, points: 1, grade: true, solution: false -->
```python
"""Check whether function F is defined and has input parameter x, and whether F(1234)=1"""

assert callable(F)  , "Sorry, F is not defined"
assert (F(x=1234) == 1)  , ""

```

<!-- grade_id: cell-238090db5fe3180a, points: 1, grade: true, solution: false -->
```python
"""Check whether F(1.54) = 0.39175318518518515"""

assert abs(F(1.54) - 0.39175318518518515) < 1e-8,  "Sorry, wrong answer"
```

<!-- grade_id: cell-b0ab0a3a503bafb4, points: 1, grade: true, solution: false -->
```python
"""Check whether in the 5th and 6th decimal place of F(2) are 9 and 2, respectively"""

assert np.mod(int(np.floor(F(2)*10**6)),100) == 92,  "Sorry, wrong answer"
```

<!-- grade_id: cell-87fa9127b01b381f, points: 5, grade: true, solution: false -->
```python
"""Check whether the answer is correct for some inputs (hidden test)"""

```

Output:
```
'Check whether the answer is correct for some inputs (hidden test)'
```

<!-- grade_id: cell-e8f511cebd625045, points: 2, grade: true, solution: false -->
```python
"""Check whether the answer is correct for some other inputs (hidden test)"""


```

Output:
```
'Check whether the answer is correct for some other inputs (hidden test)'
```


### Part c)

Create a python function `probX(a,b)` which for input values $a$ and $b$ with $a\leq b$  returns $P(a<X<b)$. You do not need to worry what if a user enters values $a$ and $b$ such that $a>b$ - you will only be tested at input values for which $a<b$.

Set the default values of `a` and `b` to be `a = float('-inf')` and `b = float('inf')` which represent $-\infty$ and $\infty$ in python. That is, `float('-inf')` is smaller than any numeric value, while `float('inf')` is larger than any numeric value.

So,  since default $b$ is $\infty$, calling `probX(a=1.2)` should return the value of $P(1.2<X<\infty)=P(X>1.2)$.

Also, since default $a$ is $-\infty$,  calling `probX(b=2.3)` should return the value of  $P(-\infty < X < 2.3) = P(X<2.3)$

In addition, `calling probX(1.2, 2.3)` should return, of course, the value of $P(1.2<X<2.3)$.

**Hint:** Recall that since $X$ is continuous, $P(X\leq b)=P(X< b)$. So, you can make use of your function `F`. Also, keep in mind that, say, value `probX(1.2, 7)` should be the same as `probX(1.2, 3)`.


**Note:** `float('inf')` is equal to `numpy.inf` from `numpy` module, which is also equal to `math.inf` from `math` module; similarly is with `float('-inf')`. The reason why you were asked to set `b=float('inf')` rather than `b=np.inf` is simply because the command `b=float('inf')` does not require importing any module, although `numpy` was imported (as `np`) in this notebook, for other purposes.


<!-- grade_id: cell-678f164042c90edc, points: 0, grade: false, solution: true -->
```python

def probX(a = float('-inf'), b = float('inf')):
    if b > 3:
        b = 3
    if a < 0:
        a = 0
    return F(b)-F(a)

```

<!-- grade_id: cell-886329ca528df713, points: 1, grade: true, solution: false -->
```python
"""Check whether probX(a,b) is defined and whether probX(-10,10)=1"""

assert callable(probX)  , "Sorry, probX is not defined"
assert abs(probX(a=-10, b=10) - 1) < 1e-5  , "Sorry, something is wrong here"

```

<!-- grade_id: cell-acdec409d93c30ca, points: 2, grade: true, solution: false -->
```python
"""Check whether probX(a=1.2)=0.744 and whether probX(b=1.2)=0.256=probX(-2.3,1.2)"""

assert abs(probX(a=1.2) - 0.744) < 1e-5  , "Sorry, wrong answer"
assert abs(probX(b=1.2) - 0.256) < 1e-5  , "Sorry, wrong answer"
assert abs(probX(a=-2.3, b=1.2) - 0.256) < 1e-5  , "Sorry, wrong answer"

```

<!-- grade_id: cell-18c94251b3bd1c59, points: 4, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```



### Part d)

Find $E[X]$, the expected value of the random variable $X$. Write your answer as

`E = <some expression>`

For example, if you think the answer is ${12 \over 34}$, set the variable `E` to be

`E = 12/34`


<!-- grade_id: cell-a319d822b97b829f, points: 0, grade: false, solution: true -->
```python

E = 7/4

```

<!-- grade_id: cell-23c101dea06dbaf2, points: 1, grade: true, solution: false -->
```python
"""Check whether E is numeric"""
assert type(E) in {int, float}  ,  "Sorry, E is not numeric"

```

<!-- grade_id: cell-8bc24128381641a6, points: 1, grade: true, solution: false -->
```python
"""Check whether in the 2nd decimal place is number 5"""

assert np.mod(int(np.floor(E*10**2)),10) == 5,  "Sorry, wrong answer"

```

<!-- grade_id: cell-7b2103fbdccddb77, points: 5, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```


### Part e)

Find $E[X^2]$ (the second order moment of $X$). Write your answer as

`EXsq = <some expression>`

For example, if you think the answer is ${12 \over 34}$, set the variable `EXsq` to be

`EXsq = 12/34`

<!-- grade_id: cell-743099ded4aba0b2, points: 0, grade: false, solution: true -->
```python

EXsq = 28/5

```

<!-- grade_id: cell-6c6093cb2205eca2, points: 1, grade: true, solution: false -->
```python
"""Check whether EXsq is numeric"""
assert type(EXsq) in {int, float}  ,  "Sorry, EXsq is not numeric"
```

<!-- grade_id: cell-d2af65bdb6468ba2, points: 1, grade: true, solution: false -->
```python
"""Check whether in the 1st decimal place is number 6"""

assert np.mod(int(np.floor(EXsq*10**1)),10) == 6,  "Sorry, wrong answer"
```

<!-- grade_id: cell-67bff0f398a3dcc0, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```


### Part f)

Find $\text{var}[X]$, the variance of the random variable $X$. Write your answer as `V`.

For example, if you think the answer is ${12 \over 34}$, set the variable `V` to be

`V = 12/34`

<br>

**Hint:** You don't need any computation here, since you can express `V` in terms of some previous answers.


<!-- grade_id: cell-69be162b9ded92d4, points: 0, grade: false, solution: true -->
```python

V = EXsq - (E**2)

```

<!-- grade_id: cell-f681891c198f811c, points: 1, grade: true, solution: false -->
```python
"""Check whether V is numeric"""
assert type(V) in {int, float}  ,  "Sorry, V is not numeric"
```

<!-- grade_id: cell-11313f249903915c, points: 1, grade: true, solution: false -->
```python
"""Check whether in the 3rd decimal place is number 7"""

assert np.mod(int(np.floor(V*10**3)),10) == 7,  "Sorry, wrong answer"
```

<!-- grade_id: cell-903f8140f3d1a9df, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""

```

Output:
```
'Check whether the answer is correct (hidden test)'
```