Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Madison Chester"
COLLABORATORS = ""
```

---

```python
##this is used for testing
import numpy as np
import types
from nose.tools import assert_equal, assert_raises
import scipy.special

```

### Binomial Coefficient $N\choose k$ in Python


The above cell imports `scipy.special`, which is a subpackage of `scipy` and contains method `comb(N,k)`, This method can be used to compute binomial coefficient ${N \choose k} = {N\cdot (N-1)\cdot \,.\,.\,.\, \cdot (N-k+1)\over k!} = {N!\over k!\,(N-k)!}$, i.e. the number of ways of choosing a subset of $k$ objects from a set of $N$ objects.

For example, if you need number of ways of choosing 3 out of 7 balls from a box, that number can be found by calling 

`scipy.special.comb(7,3)`


Feel free to use `scipy.special.comb()` in this assignment.

<br>

In addition, if you want to compute factorial $n! = n\cdot (n-1)\cdot (n-2)\cdot\, \dots \, \cdot 2\cdot 1$, you can use 

`scipy.special.factorial(n)`

<br>


### Question 1


(Recall that a regular deck of 52 cards consists of 13 kinds/denominations, each having 4 suits: <font color="red">&hearts;, &diams;</font>, &clubs; &spades;.)

Suppose you take 3 cards from a regular deck of cards, one card at a time, without replacement (i.e. without putting any card back). 


#### Part a)

What is the probability that the first, the second and the third card drawn are a Queen, a King and an Ace, respectively (i.e. in that order)? Do NOT round your answer.


Write your answer as

`ans = <some expression>`

For example, if you think the answer is `1/36` you should create variable `ans` as

`ans = 1/36`

**IMPORTANT:** Just typing `1/36`, or `answer = 1/36` is not acceptable. Also, do NOT put dot at the end (this is a python code, not a sentence in English). Your answer will be graded automatically, and the test code expects you to create a Python variable called `ans`, and the code will test the value of `ans`.


<br>


<!-- grade_id: cell-f6023e5dd46727b2, points: 0, grade: false, solution: true -->
```python

ans = (4/52) * (4/51) * (4/50)

```

<!-- grade_id: cell-66b87fb61afca9a5, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-83c5e9bb8f65e697, points: 1, grade: true, solution: false -->
```python
"""check whether the 6th and 7th decimal digits are 26"""

assert np.mod(int(np.floor(ans * 10**7)), 100) == 26

```

<!-- grade_id: cell-bff80c5cfb3c2b80, points: 2, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part b)

What is the probability that the first, the second and the third card drawn are Queen, Queen and Ace, in that order? Do NOT round your answer.

Write your answer as 

`ans = <some expression>`

<!-- grade_id: cell-47281b4aeac25572, points: 0, grade: false, solution: true -->
```python

ans = (4/52) * (3/51) * (4/50)

```

<!-- grade_id: cell-e56164e03577f41f, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-690cabf70d789b4a, points: 1, grade: true, solution: false -->
```python
"""check whether the 6th and 7th decimal digits are 19"""

assert np.mod(int(np.floor(ans * 10**7)), 100) == 19


```

<!-- grade_id: cell-02388986b9763bd5, points: 2, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 2

Consider two boxes, one containing 3 blue and 2 red marbles, the other contains 3 blue and 5 red marbles. A box is selected at random, and a marble is drawn from it at random.


#### Part a)

What is the probability that the selected marble is blue? Do NOT round your answer.

<br>

Create a variable `pB` that stores the value of this probability. That is, write your answer as

`pB = <some expression>`

<br>


<!-- grade_id: cell-51b7dd2c4839fb7d, points: 0, grade: false, solution: true -->
```python

pB = ((3/5)*(1/2)) + ((3/8)*(1/2))

```

<!-- grade_id: cell-dd7567e4b92a7299, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(pB,float)
assert (0 <= pB) & (pB <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-5ba37b4795cab141, points: 3, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part b)

If the chosen marble is blue, what is the probability that it was taken from the first box? Do NOT round your answer.

Create a variable `pFgivenB` which is equal to the desired probability.

`pFgivenB = <some expression>`


<!-- grade_id: cell-d84b5d016b3c1a43, points: 0, grade: false, solution: true -->
```python

pFgivenB = ((3/5)*(1/2))/pB

```

<!-- grade_id: cell-96d4707c591ce107, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(pFgivenB,float)
assert (0 <= pFgivenB) & (pFgivenB <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-09887c75fd5f4e23, points: 3, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 3

A blood test indicates the presence of a particular disease 97% of the time when the disease is actually present. The same test indicates the presence of the disease 1.4% of the time when the disease is not present (false positive). Two percent of the population actually has the disease. If a person gets tested and the test turns positive (suggestion the disease), what is the probability that the person does have the disease? Do NOT round your answer. 
<br>

Write your answer as

`ans = <some expression>`

<br>


<!-- grade_id: cell-81782da06f5039bd, points: 0, grade: false, solution: true -->
```python

ans = ((.97)*(.02))/(((.97)*(.02))+((.014)*(.98)))

```

<!-- grade_id: cell-40ec419bf1c4739c, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-4eb0da544a033012, points: 1, grade: true, solution: false -->
```python
"""Check whether the 6th and 7th decimal digits of the answer are 87"""

assert np.mod(int(np.floor(ans * 10**7)), 100) == 87
```

<!-- grade_id: cell-fd92e4accd3c9d4e, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 4

In a certain study a drug is found to be efficient in 67% of the people who were sick and took the drug. There are 12% of people in the study who reported minor side effects of taking the drug (soreness, headache or nausea). Also, of all the sick people for whom the drug was efficient, 8% of them had the minor side effects. 
<br>

#### Part a)

Based on this study, what is the probability that a sick person who takes the drug will both quickly recover (i.e. the drug will be efficient on them) and will have the minor side effects? Do NOT round your answer.

Store your answer as a variable `pEandS` (E - efficient, S - side effect). So, write the answer as

`pEandS = <some expression>`

<br>



<!-- grade_id: cell-767bfd46dad418e0, points: 0, grade: false, solution: true -->
```python

pEandS = (.67)*(.08) 

```

<!-- grade_id: cell-30cc2f0253090d26, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(pEandS,float)
assert (0 <= pEandS) & (pEandS <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-012f5d7c9064e9ef, points: 1, grade: true, solution: false -->
```python
"""Check whether the 3rd decimal digit is 3"""

assert np.mod(int(np.floor(pEandS * 10**3)), 10) == 3  , "Sorry, wrong answer"
```

<!-- grade_id: cell-2fe88bfa597b3915, points: 3, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part b)

Suppose a sick person takes the drug. If they experience a side effect, what is the probability the drug will be efficient on that person, i.e. the patient will quickly recover? Do NOT round your answer.

Write your answer as 

`ans = <some expression>`

<!-- grade_id: cell-a1893f225dddbbb6, points: 0, grade: false, solution: true -->
```python

ans = pEandS / .12

```

<!-- grade_id: cell-b9c1035eb0eb6b4b, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-b20a6137cd51a0e3, points: 3, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 5

(Recall that a regular deck of 52 cards consists of 13 kinds/denominations, each having 4 suits: <font color="red">&hearts;, &diams;</font>, &clubs; &spades;.)

In a poker game, you get 5 cards from a regular deck of 52 cards. 

What is the probability that you get full house (i.e. three of a kind and two of another kind) if the first two cards you've got were two aces? Do NOT round your answer. You may want to use `scipy.special` subpackage.


<br>

Write your answer as

`ans = <some expression>`

<br>


<!-- grade_id: cell-bbda68dc4b2132bd, points: 0, grade: false, solution: true -->
```python
# P(full house | first two aces) = P(full house and first two aces) / P(first two aces)
# P(full house | first two aces) = P(first two aces | full house)*P(full house) / P(first two aces)

ans = (2*((scipy.special.comb(4,2)*scipy.special.comb(12,1)*scipy.special.comb(4,3)) + (scipy.special.comb(4,3)*scipy.special.comb(12,1)*scipy.special.comb(4,2))) / scipy.special.comb(52,5)) / ((scipy.special.comb(4,2)*scipy.special.comb(50,3)) / scipy.special.comb(52,5)) 

```

<!-- grade_id: cell-58245af72aab38ed, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-f6800c7e4aea1822, points: 1, grade: true, solution: false -->
```python
"""Check whether the 6th and 7th decimal digits of the answer are 59"""

assert np.mod(int(np.floor(ans * 10**7)), 100) == 59
```

<!-- grade_id: cell-0445ee461dd2e396, points: 6, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

### Question 6

Members of a consulting firm rent cars from 3 rental agencies: 60% from agency $A_1$, 30% from agency $A_2$, and 10% from $A_3$. Suppose 9% of the cars from agency $A_1$ need oil change, 20% of cars from $A_2$ need oil change and 6% of cars from $A_3$ need an oil change.

#### Part a)

What is the probability that a rental car delivered to the firm will need an oil change? Do NOT round your answer.

<br>

As an answer to this question, create a variable `pC` whose value is the desired probability that the delivered car needs an oil change. So, your answer should be of the form

`pC = <some expression>`

<br>




<!-- grade_id: cell-a10b3d248b7ec37b, points: 0, grade: false, solution: true -->
```python

pC = (.09)*(.60) + (.20)*(.30) + (.06)*(.10)

```

<!-- grade_id: cell-878482512d275b36, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(pC,float)
assert (0 <= pC) & (pC <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-02b027b15987404f, points: 3, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```

#### Part b)

If a car delivered to the consulting firm needs oil change, what is the probability that it came from rental agency $A_2$? Do NOT round your answer.

Write your answeer as a variable `ans`, i.e.

`ans = <some expression>`


<!-- grade_id: cell-7f94ea968a38fe73, points: 0, grade: false, solution: true -->
```python

ans = ((.20)*(.30))/pC

```

<!-- grade_id: cell-868e5e1697c1f48b, points: 1, grade: true, solution: false -->
```python
"""Check whether the answer is numeric and from the interval [0,1] - since it's a probability"""

assert isinstance(ans,float)
assert (0 <= ans) & (ans <= 1)  , "Sorry, the answer should be a probability, which means a number from [0,1]"

```

<!-- grade_id: cell-25951e4076ac2771, points: 4, grade: true, solution: false -->
```python
"""Check whether the answer is correct (hidden test)"""


```

Output:
```
'Check whether the answer is correct (hidden test)'
```